---
title: 'Caching, logging, and query plans in Spark SQL'
description: ""
---

## Capstone

```yaml
type: NormalExercise
key: 6b93918e78
xp: 100
```



`@instructions`


`@hint`


`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
